Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

```python
# This line configures matplotlib to show figures embedded in the notebook, 
# instead of opening a new window for each figure.
%matplotlib inline
# change the value after "width:" for a nice visualization in you computer
from IPython.core.display import display, HTML
display(HTML("<style>.container { width:70% !important; }</style>"))
```

## Attività di laboratorio/seminario

### Python per il calcolo scientifico

#### Esercitazione 1

**Docente**: Lucas Omar Müller

**Email**: lucas.muller@unitn.it

<span style="color:red">Informazione nel [sito Moodle del corso](https://didatticaonline.unitn.it/dol/course/view.php?id=28393)!</span>

### Esercizio 1:

Scrivere una funzione che:

* prende come argomento una lista
* verifica che tutti i suoi elementi siano degli interi
* verifica che non ci siano elementi ripetuti
* ordina gli elementi da minore a maggiore (senza usare la funzione `sort` della classe `list`!)
* restituisce un `tuple` dove 
    * il primo elemento è una variabile `None` se la funzione non identifica non interi tra gli elementi della lista o valori ripetuti, oppure l'eccezione sollevata
    * il secondo elemento è l'indice del primo elemento di tipo non intero nella lista, oppure il tuple dei primi due elementi ripetuti, oppure la lista ordinata

<!-- grade_id: cell-c67211bf6b2b2343, points: 0, grade: false, solution: true -->
```python
def ordinaInteri(listaIn):
    """ 
    verifica che tutti gli elementi di listaIn siano interi. Se questo si verifica, 
    restituisce una lista con gli elmenti di listaIn ordinati in ordine crescente,
    altrimenti genera una eccezione e restituisce il tipo di eccezione è 
    l'indice del primo elemento che non è di tipo intero
    
    Argomenti in entrata:
    - listaIn: lista
    Argomenti in uscita:
    - status: None se tutti gli elementi di listaIn sono interi e non vi sono elementi ripetuti, altrimenti eccezione
    - result: 
        - se status non è None:
                  - indice del primo elemento di tipo non intero
                  - tuple con gli indici dei primi due elementi ripetuti,
        - altrimenti, lista con elementi di listaIn ordinati da minore a maggiore
    """ 
    # YOUR CODE HERE
    raise NotImplementedError()
    
    
    ### TRACCIA
    
    # initializzare "status"
    # status = 
    # assegnare alla variabile "n" la lunghezza della lista (funzione "len")
    # n = 
    # verificare se ci sono non interi
    # loop "for" su tutti gli elementi della lista:
    #      A) faccio un "try":
    #      se l i-essimo elemento non è intero, allora uso "raise" per 
    #      un "TypeError" exception (a lezione abbiamo visto il "Exception"), tipo:
    #            raise TypeError("Questo non è un intero")
    #      B) se ottengo un exception, allora faccio "except Exception as e":
    #      e li dentro assegno "i" a "result" e "e" a "status"
    #      C) se abbiamo incontrato valori non interi dobbiamo usciere dal loop, per esempio usando "break"
    
    # se tutti interi, verificare che non ci siano valori ripetuti
    # loop "for" su tutti gli elementi della lista (in "i"):
    #      A) assegno a "val" il valore del i-essimo elmento della lista
    #      B) altro loop sugli elementi della lista (in "j"):
    #            B.1) se "i" diverso da "j"
    #            B.2) assegno a "valOther" il valore del j-essimo elmento della lista
    #            B.3) faccio un "try "except", in questo caso l'exception viene
    #                 causata da valori uguali, usiamo quindi (dentro il "try":)
    #                     "raise ValueError(Valori ripetuti)"          
    #                 dentro l'except Exception as e" assegnamo:
    #                     un tuple on "i" e "j" per "result"
    #                     "e" per "status" 
    #            B.4) se abbiamo incontrato valori ripetuti dobbiamo usciere dal loop annidato
    #      C) se abbiamo incontrato valori ripetuti dobbiamo usciere dal loop
    # se tutto a posto allora possiamo ordinare
    # se "status" è ancora "None"
    #   Inizializziamo "result" come un alista vuota
    #   Appendiamo il primo elmento di listaIn
    #   loop for dal secondo all'ultimo elemento di listaIn
    #       A) assegnamo a "newValue" l'i-essimo elmento della lista
    #       B) se "newValue" minore del primo elemento di "result" allora lo inseriamo come prim
    #       C) se "newValue" maggiore che l'ultimo elmento di "result" allora lo mettiamo come ultimo
    #       D) altrimenti dobbiamo inserire newValue tra due elementi di result
    #          quindi facciamo un loop con "j" che va da "0" a "len(result)-1"
    #              se "newValue" è maggiore che result[j] e minore che "result[j+1]"
    #              lo inseriamo dopo "result[j]" e fermiamo il ciclo con break
     
    
    return status, result
```

```python
lista = [0,5,2.3,10]
#lista = [0,10.1,2,10.1]
#lista = [0,5,2,5]
#lista = [6,3,1,5,2,4]
#lista = [0,5,2,10]
status, result = ordinaInteri(lista)
print(status)
print(type(status))
print(result)
```

<!-- grade_id: cell-5d9eddd8c58db04e, points: 1, grade: true, solution: false -->
```python
lista = [0,5,2.3,10]
status, result = ordinaInteri(lista)
print(status)
print(type(status))
print(result)
assert(status!=None)
assert(isinstance(status,TypeError))
assert(result==2)

```

<!-- grade_id: cell-5be19b008c1173bc, points: 1, grade: true, solution: false -->
```python
lista = [0,5,2,5]
status, result = ordinaInteri(lista)
print(status)
print(type(status))
print(result)
assert(status!=None)
assert(isinstance(status,ValueError))
assert(result==(1,3))
```

<!-- grade_id: cell-768dd3632441876c, points: 2, grade: true, solution: false -->
```python
lista = [6,3,1,5,2,4]
status, result = ordinaInteri(lista)
print(status)
print(type(status))
print(result)
assert(status==None)
assert(result==[1,2,3,4,5,6])
```

```python

```

```python

```